---
title: 'Setup RStudio'
description: 'Chapter description goes here.'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 7683200cc3
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}
success_msg("That looked really good!", TRUE)
```
